![CC-BY-SA](https://mirrors.creativecommons.org/presskit/buttons/88x31/svg/by-sa.svg)


This notebook was created by [Bernardo Freitas Paulo da Costa](http://www.im.ufrj.br/bernardofpc),
and Pedro Angelo Medeiros Fonini
and is licensed under Creative Commons BY-SA.

Antes de enviar este Teste, verifique que tudo está funcionando como esperado.
Por exemplo, **rode o código inteiro, do zero**.
Para isso, vá no menu, escolha _Kernel_, depois _Restart & Run All_.

Verifique, também, que você respondeu todas as questões:
* as questões de código têm `YOUR CODE HERE` (e você pode apagar o `raise NotImplemented` ao incluir sua resposta)
* as questões discursivas têm "YOUR ANSWER HERE".

---

```python
import numpy as np
import matplotlib.pyplot as plt
```

```python
def veryclose(v1,v2):
    return np.allclose(v1, v2, atol=1e-13, rtol=1e-13)
```

<!-- grade_id: cell-53850929017d728c, points: 0, grade: false, solution: true -->
```python
def lagrange(xs, ys):
    assert len(xs) == len(ys)
    # YOUR CODE HERE
    raise NotImplementedError()
```

# Parte 1: Interpolação e Derivação

## Questão 1: Derivando a interpolação

Escreva uma função que, dados pontos $(x_i, y_i)$,
retorna os coeficientes da **derivada** do polinômio interpolador de Lagrange.

<!-- grade_id: cell-3a32d6699e64c7c7, points: 0, grade: false, solution: true -->
```python
def lagrange_coefs_linha(xs, ys):
    assert len(xs) == len(ys)
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-2064385cebf54654, points: 1, grade: true, solution: false -->
```python
assert veryclose(lagrange_coefs_linha([0,7], [4,6]), [2/7])
```

<!-- grade_id: cell-e2f45275c1c8b937, points: 1, grade: true, solution: false -->
```python
assert veryclose(lagrange_coefs_linha([0,1,2], [3,4,6]), [0.5, 1.])
```

<!-- grade_id: cell-1a1692de98986fff, points: 1, grade: true, solution: false -->
```python
assert veryclose(lagrange_coefs_linha([0,1,2,-1], [3,4,6,7]), [-5/6, 5, -2])
```

<!-- grade_id: cell-a6473543c3af7ab4, points: 1, grade: true, solution: false -->
```python
assert veryclose(lagrange_coefs_linha([4,5,9,10], [0,3,0,3]), [35, -10.5, 0.75])
```

Agora, use a `lagrande_coefs_linha` para escrever a função `lagrange_linha(xs,ys)`,
que retorna uma função que calcula este polinômio.

<!-- grade_id: cell-2e85b2f87e3f19b5, points: 0, grade: false, solution: true -->
```python
def lagrange_linha(xs, ys):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-b3d70fc87bcece02, points: 1, grade: true, solution: false -->
```python
f = lagrange_linha([0,1,2], [3,2,3])
assert veryclose(f([0,1,2]), [-2, 0, 2])
```

<!-- grade_id: cell-308629e44ccc6379, points: 1, grade: true, solution: false -->
```python
f = lagrange_linha([3,4,5,6,7], [1,2,1,4,3])
assert veryclose(f([3,4,5,6,7])*3, [45/2, -13/2, 7/2, 21/2, -55/2])
```

## Questão 2: Gráficos

Faça um gráfico do **erro absoluto** cometido pelos polinômios interpoladores de grau 4 e 9
para a função seno no intervalo $[-1,4]$.

<!-- grade_id: cell-daf55da8b3da0ece, points: 0, grade: true, solution: true -->
```python
I = [-1, 4]
ts   = np.linspace(I[0], I[1], num=300)
# YOUR CODE HERE
raise NotImplementedError()
ax = plt.gca()
plt.show()
```

<!-- grade_id: cell-a9a3290541b9be90, points: 1, grade: true, solution: false -->
```python
assert ax.title.get_text() != ""
assert len(ax.lines) == 2
assert len(ax.legend().texts) == 2
```

<!-- grade_id: cell-078d4f6225002105, points: 2, grade: true, solution: false -->
```python
for l in ax.lines:
    xdata = l.get_xdata()
    ydata = l.get_ydata()

    assert min(xdata) == -1
    assert max(xdata) ==  4

    assert 1e-5 <= max(ydata) <=  0.035
    assert min(ydata) >=  0
```

```python
ax = None
```

Agora, faça os gráficos dos erros correspondentes às derivadas,
ou seja, o quanto a derivada da interpolação está longe da derivada real da função.

<!-- grade_id: cell-170d9d055782a56d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
ax = plt.gca()
plt.show()
```

<!-- grade_id: cell-560fdb129b228274, points: 1, grade: true, solution: false -->
```python
assert ax.title.get_text() != ""
assert len(ax.lines) == 2
assert len(ax.legend().texts) == 2
```

<!-- grade_id: cell-a6a0e2b425db3401, points: 1, grade: true, solution: false -->
```python
for l in ax.lines:
    xdata = l.get_xdata()
    ydata = l.get_ydata()

    assert min(xdata) == -1
    assert max(xdata) ==  4

    assert max(ydata) <=  0.25
    assert min(ydata) >=  0
```

```python
ax = None
```

Enfim, faça dois gráficos dos erros absolutos para o seno e sua derivada,
à esquerda a partir do polinômio interpolador de grau 4, e à direita a partir do de grau 9.

<!-- grade_id: cell-fca661da50aa6997, points: 0, grade: false, solution: true -->
```python
fig, [ax1, ax2] = plt.subplots(ncols=2, figsize=(12,4))

# YOUR CODE HERE
raise NotImplementedError()
plt.show()
```

<!-- grade_id: cell-60b75d4fa6952413, points: 1, grade: true, solution: false -->
```python
assert len(ax1.lines) == 2

for l in ax1.lines:
    xdata = l.get_xdata()
    ydata = l.get_ydata()

    assert min(xdata) == -1
    assert max(xdata) ==  4

    assert max(ydata) <=  0.25
```

<!-- grade_id: cell-42196daa797f1c72, points: 1, grade: true, solution: false -->
```python
assert len(ax2.lines) == 2

for l in ax2.lines:
    xdata = l.get_xdata()
    ydata = l.get_ydata()

    assert min(xdata) == -1
    assert max(xdata) ==  4

    assert max(ydata) <=  1e-3
```

```python
ax1 = None
ax2 = None
```

- O que você observa nestes dois gráficos, quanto ao erro da derivada?
- Como esses erros se comparam com os erros de estimativa da função propriamente dita?
- Como você explicaria este(s) fenômeno(s)?

<!-- grade_id: cell-1b3f9430479f4aea, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

## Questão 3: Interpolando a derivada

Escreva uma função que, dados os pontos $(x_i, y_i)$,
retorna o polinômio interpolador correspondente à estimativa da derivada central gerada por estes pontos.
(você pode supor que as abscissas $x_i$ estão em ordem).

Dica: quais são as abscissas das derivadas centrais?

<!-- grade_id: cell-2aad0c19440b5e2c, points: 0, grade: false, solution: true -->
```python
def lagrange_central(xs, ys):
    assert len(xs) == len(ys)
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-599c5e550f924ae4, points: 2, grade: true, solution: false -->
```python
f = lagrange_central([0,1,2], [3,2,3])
assert veryclose(f([0,1,2]), [-2, 0 , 2])
```

<!-- grade_id: cell-556845aaaa91bc13, points: 2, grade: true, solution: false -->
```python
f = lagrange_central([0,1,2,3], [3,2,3,2])
assert np.allclose(f([.3, 1.5, 2.8]), [-1.88, 1, -2.38])
```

## Questão 4: Mais gráficos

Agora, faça os gráficos de erros absolutos para a derivada, usando polinômios de grau 4 e 9, novamente.

<!-- grade_id: cell-151cb87420c1ab0c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
ax = plt.gca()
plt.show()
```

<!-- grade_id: cell-11639a33fb4f31e1, points: 1, grade: true, solution: false -->
```python
assert ax.title.get_text() != ""
assert len(ax.lines) == 2
assert len(ax.legend().texts) == 2
```

<!-- grade_id: cell-f0f84feedcd6ec19, points: 1, grade: true, solution: false -->
```python
for l in ax.lines:
    xdata = l.get_xdata()
    ydata = l.get_ydata()

    assert min(xdata) == -1
    assert max(xdata) ==  4

    assert max(ydata) <=  0.35
    assert min(ydata) >=  0
```

```python
ax = None
```

Novamente, faça lado a lado o gráfico para graus 4 e 9,
só que agora para o erro cometido pelos dois métodos diferentes de estimativa da derivada.

<!-- grade_id: cell-e7dfb9b8b4aa2c88, points: 0, grade: false, solution: true -->
```python
fig, [ax1, ax2] = plt.subplots(ncols=2, figsize=(12,4))

# YOUR CODE HERE
raise NotImplementedError()
plt.show()
```

<!-- grade_id: cell-7c380873132d49d1, points: 2, grade: true, solution: false -->
```python
assert len(ax1.lines) == 2
assert len(ax1.legend().texts) == 2

for l in ax1.lines:
    xdata = l.get_xdata()
    ydata = l.get_ydata()

    assert min(xdata) == -1
    assert max(xdata) ==  4

    assert 0.1 <= max(ydata) <=  0.35
    assert min(ydata) >=  0
```

<!-- grade_id: cell-410d47b490797b7c, points: 2, grade: true, solution: false -->
```python
assert len(ax2.lines) == 2
assert len(ax2.legend().texts) == 2

for l in ax2.lines:
    xdata = l.get_xdata()
    ydata = l.get_ydata()

    assert min(xdata) == -1
    assert max(xdata) ==  4

    assert 4e-4 <= max(ydata) <=  0.02
    assert min(ydata) >=  0
```

```python
ax1 = None
ax2 = None
```

- Qual método produziu menores erros?
- Qual método melhor aproveitou a maior quantidade de pontos disponíveis?

Como você explicaria estas conclusões?

<!-- grade_id: cell-862109910b6ff5e6, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

# Parte 2: Integração e regressão

Imagine que temos alguns dados $(x_i, z_i)$, e queremos calcular a integral de $f$ no intervalo $I$.
Entretanto, os valores $z_i$ não são iguais a $y_i = f(x_i)$, mas contém um ruído $\varepsilon_i$,
que supomos pequeno, mas não nulo.
O objetivo desta parte da prova é estudar como os ruídos em $z_i$ se transformam em erros na estimativa da integral.

# Questão 5: Simpson

Dê abaixo a regra de Simpson para integrar funções dividindo um intervalo em $n$ partes iguais.
Este método é suficientemente preciso para ser usado como "referência".

<!-- grade_id: cell-e0ced51693d7b29f, points: 0, grade: false, solution: true -->
```python
def simpson(f, a, b, n=1000):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-15f383c0f2c6950c, points: 1, grade: true, solution: false -->
```python
assert np.isclose(simpson(np.sin, 0, np.pi, n=50), 2)
assert np.isclose(simpson(np.sin, 0, np.pi, n=500), 2, atol=1e-12, rtol=1e-12)
```

<!-- grade_id: cell-5ed1982a348645a5, points: 1, grade: true, solution: false -->
```python
def f(x): return np.exp(x/4) - 0.7*np.sin(2.3*x)
assert np.isclose(simpson(f, -4, 10, n=50),   47.3929830436714)
assert  veryclose(simpson(f, -4, 10, n=2000), 47.3929830436714)
```

# Questão 6: Comparando integrais

Considere a função $f$ abaixo, e calcule o polinômio interpolador de $f$ nos 21 pontos $(x_i, y_i)$ dados.
Use-o para estimar a integral de $f$ no intervalo $[-2,3]$.

```python
def f(t):
    return np.exp(t/8)*np.sin(2*t) + np.exp(-t/9)*np.cos(np.pi*t)

npts = 21
xs = np.linspace(-2,3, num=npts)
ys = f(xs)
```

<!-- grade_id: cell-687c46785e52e513, points: 0, grade: false, solution: true -->
```python
def poly_int(coefs, a, b):
    """Calcula a integral do polinômio com coeficientes `coefs` dados em ordem crescente de grau, no intervalo [a,b]."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-36f17fff49ec383b, points: 0, grade: false, solution: true -->
```python
### Dê sua resposta na forma I = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-bcf23ab333f938fb, points: 2, grade: true, solution: false -->
```python
# Teste aqui...
```

Agora, faça o mesmo com os pontos $(x_i, z_i)$, que estão sujeitos a um erro de "magnitude média" $0.1$.
Qual o erro de integração?

```python
np.random.seed(2)
A = 0.1
zs = ys + np.random.randn(len(ys))*A
```

<!-- grade_id: cell-680a0b02dff25279, points: 0, grade: false, solution: true -->
```python
### Dê sua resposta na forma Iz = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c5243ec607befa4a, points: 2, grade: true, solution: false -->
```python
# Teste aqui...
```

## Questão 7: Comparando com o método de Simpson

Quantos intervalos são necessários para que o método de Simpson
tenha um erro menor do que usar o polinômio interpolador com os dados **corretos**?

<!-- grade_id: cell-4f6e3653c2058f03, points: 0, grade: false, solution: true -->
```python
# Dê sua resposta na forma n_int = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3a647834467f7de7, points: 2, grade: true, solution: false -->
```python
# Teste aqui...
```

Agora, supondo que você tem acesso **apenas aos pontos errados** $(x_i, z_i)$,
use-os como se fossem os valores do método de Simpson no intervalo $[-2,3]$.
Qual a estimativa da integral que você obtém?

<!-- grade_id: cell-9ab4136f6b3f5324, points: 0, grade: false, solution: true -->
```python
# Dê sua resposta na forma Iz_simp = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-043f2c8fc4ddfcc6, points: 2, grade: true, solution: false -->
```python
# Teste aqui...
```

## Questão 8: Gráficos

Faça o gráfico dos erros de interpolação para a função $f$,
tanto para o polinômio interpolador com dados "certos" como para o polinômio gerado pelos dados com ruído.

<!-- grade_id: cell-9d22a864a621dd46, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
ax = plt.gca()
plt.show()
```

<!-- grade_id: cell-14e0e08f744977bb, points: 2, grade: true, solution: false -->
```python
assert len(ax.lines) == 2
assert len(ax.legend().texts) == 2

for l in ax.lines:
    xdata = l.get_xdata()
    ydata = l.get_ydata()

    assert min(xdata) == -2
    assert max(xdata) ==  3

    assert max(ydata) <=  40
    assert min(ydata) >=  0
```

```python
ax = None
```

Como o gráfico acima ajuda a explicar o erro de integração ao usar pontos com ruído?

<!-- grade_id: cell-606e115318f8c968, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

Faça o gráfico do erro de interpolação correspondente a usar o método de Simpson com os pontos $(x_i, z_i)$.

Atenção: Observe que o método de Simpson irá formar polinômios diferentes em intervalos diferentes.

<!-- grade_id: cell-486b0e6879dfb7f3, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
ax = plt.gca()
plt.show()
```

<!-- grade_id: cell-b7ffa334d0cbdc4e, points: 1, grade: true, solution: false -->
```python
for l in ax.lines:
    xdata = l.get_xdata()
    ydata = l.get_ydata()

    assert min(xdata) >= -2
    assert max(xdata) <=  3

    assert 0.01 <= max(ydata) <=  1
    assert min(ydata) >=  0
```

```python
ax = None
```

## Questão 9: Regressão

Agora, vamos usar 21 pontos para ajustar um polinômio de grau 10.
Qual a estimativa da integral tanto para os pontos "certos" como para os pontos "com erro"?

<!-- grade_id: cell-823a1826f2547c5a, points: 0, grade: false, solution: true -->
```python
def regres_coefs(xs, ys, deg):
    """Calcula os coeficientes da regressão polinomial. """
    assert len(xs) == len(ys)
    assert len(xs) > deg
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-14624e0adb5ea6ac, points: 0, grade: false, solution: true -->
```python
### Dê a resposta na forma
#   I_certos = ...
#   I_ruidos = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-30b58fca7fc69423, points: 2, grade: true, solution: false -->
```python
assert np.isclose(I_certos, -0.96, atol=1e-1, rtol=1e-1)
assert np.isclose(I_ruidos, -1.06, atol=1e-1, rtol=1e-1)
```

Mantendo o uso dos 21 pontos, varie o grau de regressão entre 0 e 18,
e faça um gráfico do erro de itegração tanto para o uso dos pontos corretos como dos pontos com ruído.

<!-- grade_id: cell-8303b3b0d0a75d6b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
ax = plt.gca()
plt.show()
```

<!-- grade_id: cell-72a5b75b30b478f0, points: 2, grade: true, solution: false -->
```python
assert len(ax.lines) == 2
assert len(ax.legend().texts) == 2

for l in ax.lines:
    xdata = l.get_xdata()
    ydata = l.get_ydata()

    assert min(xdata) ==  0
    assert max(xdata) == 18

    assert max(ydata) <=  200
    assert min(ydata) >=  0
```

- Como o erro se comporta ao aumentar o grau, para cada método?
- Como interpretar o erro da integral ao usar uma regressão de grau zero?

<!-- grade_id: cell-ad148db5ad08fb98, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

Agora, divida os 21 pontos em dois blocos de 11 (o do meio aparece nos dois!),
e ajuste dois polinômios, um para cada lado do intervalo de integração.
Novamente, faça o gráfico comparando o erro cometido ao usar os pontos certos e os com ruído,
em função do grau de interpolação.

<!-- grade_id: cell-408a2145460ab615, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
ax = plt.gca()
plt.show()
```

<!-- grade_id: cell-6e8fbcd7332260ba, points: 2, grade: true, solution: false -->
```python
assert len(ax.lines) == 2
assert len(ax.legend().texts) == 2

for l in ax.lines:
    xdata = l.get_xdata()
    ydata = l.get_ydata()

    assert min(xdata) ==  0
    assert max(xdata) <= 10

    assert max(ydata) <=  2
    assert min(ydata) >=  0
```

```python
ax = None
```

## Questão 10: "Amostragem"

De maneira mais geral, o ruído que adicionamos a `ys` para formar o vetor `zs` poderia ter sido outro,
por exemplo variando a semente do gerador aleatório, sem mudar a "amplitude".

Faça, agora, o gráfico do erro de integração, para regressões de grau 0 a 18,
usando 6 amostras diferentes para `zs`,
e compare com o decaimento dos erros quando não há ruídos.

<!-- grade_id: cell-c0feab8e5b2558f5, points: 2, grade: true, solution: true -->
```python
np.random.seed(123)

# YOUR CODE HERE
raise NotImplementedError()

plt.legend(bbox_to_anchor=(1.0, 0.5), loc="center left")
ax = plt.gca()
plt.show()
```

- O que você observa aqui?
- Qual ordem de regressão você escolheria, sem saber qual vai ser a realização do ruído?
- Porquê?

<!-- grade_id: cell-f367fb6d393f2255, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE